# Python Fundamentals 

The purpose of this notebook is to determine how familiar you are with Python. The grades associated with this notebook are trivial; it is mainly used to help me understand where the classes' skills are at.

Please do not consult other persons or resources, though you may reference the Python documentation [here](https://docs.python.org/3/). You have 70 minutes to complete as many tasks as you can. Please push your changes to GitHub before the end of the period even if you are unable to complete the entire assessment. 

## Short Answer
Try and go through the short answer quickly to get to the programming, if you don't know an answer tell me what you do know about the topic.

In your own words, what is recursion and when would you use it? When would you not use it?

<!-- grade_id: cell-4cdf2f8a5ec87932, points: 3, grade: true, solution: true -->


```python
def fact(n):
    total = 1
    for i in range(n, 0, -1):
        total *= i
    return total

fact(4)

def fact_recursion(n):
    if n == 1:
        return 1
    return n * fact_recursion(n-1)

fact_recursion(4)
```

Output:
```
24
```

In your own words, what is the difference between an object and a class in Python?

<!-- grade_id: cell-78419b27357c3b3f, points: 2, grade: true, solution: true -->


In your own words, what is the difference between a dictionary and a list in Python?

<!-- grade_id: cell-cad0eda710a4b3c4, points: 2, grade: true, solution: true -->


In your own words, what is a Python generator?

<!-- grade_id: cell-5f2241eb32b9cf7f, points: 2, grade: true, solution: true -->


In your own words, what does `%` do in Python?

<!-- grade_id: cell-9dfdd8dc825e6334, points: 2, grade: true, solution: true -->


In your own words, what is a `lambda` in Python?

<!-- grade_id: cell-6db62aa4184b1be2, points: 2, grade: true, solution: true -->
f = lambda x: x+1

```python
f = lambda x: x+1
f(2)
```

Output:
```
3
```

## Problem 1
A palindrome is a string which, when the characters is reversed, is equivalent to the original string. For instance, the string `radar` and `mom` are palindromes. Write a function called `check_palindrome(s)` which takes a single parameter, and returns `True` if it is a string and palindrome, `False` if it is a string but not a palindrome, and `None` if it is not a string.

<!-- grade_id: cell-977a56583b67539d, points: 0, grade: false, solution: true -->
```python
def check_palindrome(s):
    print(s[::-1])
    return s == s[::-1]


```

Output:
```
hsif
False
```

<!-- grade_id: cell-bb6fc46704d53edc, points: 4, grade: true, solution: false -->
```python
### DO NOT MODIFY THIS CELL
### BEGIN HIDDEN TESTS
assert check_palindrome('dad') == True
assert check_palindrome(5) is None
assert check_palindrome('fish') == False
assert check_palindrome('amanaplanacanalpanama') == True
assert check_palindrome('x') == True
### END HIDDEN TESTS
```

Output:
```
dad
5
<ipython-input-14-8cada89cf1c2> in <module>
      2 ### BEGIN HIDDEN TESTS
      3 assert check_palindrome('dad') == True
----> 4 assert check_palindrome(5) is None
      5 assert check_palindrome('fish') == False
      6 assert check_palindrome('amanaplanacanalpanama') == True

<ipython-input-13-87832c019a29> in check_palindrome(s)
      1 def check_palindrome(s):
      2     print(s)
----> 3     return s == s[::-1]

TypeError: 'int' object is not subscriptable
```

## Problem 2
A number can be a palindrome too. For instance, `151` is a palindrome, as is `12677621`. Modify the function `check_palindrome(s)` so that `True` is returned if `s` is a palindrome (regardless of whether `s` is an `int` or `string`) and `False` otherwise.

<!-- grade_id: cell-b5629624f441d769, points: 0, grade: false, solution: true -->
```python
def check_palindrome(s):
    s = str(s)
    return s == s[::-1]
```

<!-- grade_id: cell-6cce999c3eb24cd7, points: 3, grade: true, solution: false -->
```python
### DO NOT MODIFY THIS CELL
### BEGIN HIDDEN TESTS
assert check_palindrome('bob') == True
assert check_palindrome(5) == True
assert check_palindrome('fish') == False
assert check_palindrome(12688621) == True
assert check_palindrome(12359238) == False
### END HIDDEN TESTS
```

## Problem 3
Write a function called `add_palindrome(a,b)` that takes two parameters, which you may assume are of type `int`, adds them together only if they are both `palindromes`, and otherwise throws an exception. For instance, `add_palindrome(202,313)` should return the number `515`, but `add_palindrome(202,223)` should thrown an exception.

<!-- grade_id: cell-8b8404fa08935d72, points: 0, grade: false, solution: true -->
```python
def add_palindrome(a, b):
    if check_palindrome(a) and check_palindrome(b):
        return a + b
    else:
        raise Exception()

        
```

<!-- grade_id: cell-cccadb4c9d7d9216, points: 4, grade: true, solution: false -->
```python
### DO NOT MODIFY THIS CELL
### BEGIN HIDDEN TESTS
import unittest
tc = unittest.TestCase()
assert add_palindrome(303,121) == 424
tc.assertRaises(BaseException, lambda: add_palindrome(202,223))
### END HIDDEN TESTS
```

## Problem 4
Modify `check_palindrome(s)` so that if `s` is of type `int` it will also work will zero padded numbers (e.g. `10` is a palindrome, as is `200`, since they could be expressed as `010` and `00200` in string form).

<!-- grade_id: cell-b34e4610e817827b, points: 0, grade: false, solution: true -->
```python
def check_palindrome(s):
    s = str(s).strip('0')
    return s == s[::-1]
```

<!-- grade_id: cell-9c5b048518a905eb, points: 4, grade: true, solution: false -->
```python
### DO NOT MODIFY THIS CELL
### BEGIN HIDDEN TESTS
assert check_palindrome(20) == True
assert check_palindrome(300) == True
assert check_palindrome(301) == False
assert check_palindrome(4000) == True
assert check_palindrome(4100) == False
assert check_palindrome(40040) == False
assert check_palindrome(12688621) == True
assert check_palindrome(12359238) == False
### END HIDDEN TESTS
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-24-4ff77a1e4a88> in <module>
      6 assert check_palindrome(4000) == True
      7 assert check_palindrome(4100) == False
----> 8 assert check_palindrome(40040) == False
      9 assert check_palindrome(12688621) == True
     10 assert check_palindrome(12359238) == False

AssertionError:
```

## Problem 5
Given the string in the format `'visa,29022345,5.09:visa,888223442,10.10:mastercard,9993345522,6.00'` which is a fake list of credit card transactions where the card type is followed by the card number and the cost billed, write a function `add_visa_costs(s)` which adds up all costs in the string which are associated with visa cards (in this case the return value should be `15.19`.

<!-- grade_id: cell-2a0cff773a2b4c5e, points: 0, grade: false, solution: true -->
```python
def add_visa_costs(s):


```

<!-- grade_id: cell-fab2d6194e7ef2ac, points: 5, grade: true, solution: false -->
```python
### DO NOT MODIFY THIS CELL
### BEGIN HIDDEN TESTS
test_str = 'mastercard,29022325,3.09:visa,888223242,8.52:mastercard,9893345522,2.00:visa,7883345522,9.01'
assert add_visa_costs(test_str) == 17.53
### END HIDDEN TESTS
```

## Problem 6
Write a function `apply(f, l)` which takes a function `f` and executes this function against all of the items in `list` `l`. Your `apply()` function should return a new list of the items from `l`. To test this, define `f` to just increment a number by one, and set `l` to be a list of numbers. So `apply(f,[1,2,3])` should return `[2,3,4]`.

<!-- grade_id: cell-63baea23e1688b28, points: 0, grade: false, solution: true -->
```python
def apply(f, l):
    return [f(item) for item in l]

f = lambda x: x+ 1
print(apply(f, [1, 2, 3]))
```

Output:
```
[2, 3, 4]
```

<!-- grade_id: cell-592c5f3f1d51766e, points: 3, grade: true, solution: false -->
```python
### DO NOT MODIFY THIS CELL
### BEGIN HIDDEN TESTS
test_l = [0, 10, 100]
assert apply(lambda x: x + 3, test_l) == [3, 13, 103]
assert apply(lambda x: x ** 2, test_l) == [0, 100, 10000]
### END HIDDEN TESTS
```

Output:
```
[0, 10, 100]
<ipython-input-30-aec70a144fc9> in <listcomp>(.0)
      1 def apply(f, l):
      2     print(l)
----> 3     return [f(l) for item in l]
      4 
      5 f = lambda x: x+ 1

<ipython-input-31-99a753996569> in <lambda>(x)
      2 ### BEGIN HIDDEN TESTS
      3 test_l = [0, 10, 100]
----> 4 assert apply(lambda x: x + 3, test_l) == [3, 13, 103]
      5 assert apply(lambda x: x ** 2, test_l) == [0, 100, 10000]
      6 ### END HIDDEN TESTS

TypeError: can only concatenate list (not "int") to list
```

## Problem 7
Imagine you work for a company who writes their usernames in the format of 'aa11' where the first two characters are lower case letters and the last two are numbers. Write a function `generate_usernames()` which returns a *generator* of all possible usernames.

<!-- grade_id: cell-fd50531bf0e25026, points: 0, grade: false, solution: true -->
```python
def generate_usernames():
    letters = 'abcdefghijiklmnopqrstuvwxyz'
    numbers = '0123456789'
    
    for x in letters:
        for y in letters:
            for z in numbers:
                for w in numbers:
                    yield x+y+z+w
    
for x in generate_usernames():
    print(x)


```

Output:
```
aa00
aa01
aa02
aa03
aa04
aa05
aa06
aa07
aa08
aa09
aa10
aa11
aa12
aa13
aa14
aa15
aa16
aa17
aa18
aa19
aa20
aa21
aa22
aa23
aa24
aa25
aa26
aa27
aa28
aa29
aa30
aa31
aa32
aa33
aa34
aa35
aa36
aa37
aa38
aa39
aa40
aa41
aa42
aa43
aa44
aa45
aa46
aa47
aa48
aa49
aa50
aa51
aa52
aa53
aa54
aa55
aa56
aa57
aa58
aa59
aa60
aa61
aa62
aa63
aa64
aa65
aa66
aa67
aa68
aa69
aa70
aa71
aa72
aa73
aa74
aa75
aa76
aa77
aa78
aa79
aa80
aa81
aa82
aa83
aa84
aa85
aa86
aa87
aa88
aa89
aa90
aa91
aa92
aa93
aa94
aa95
aa96
aa97
aa98
aa99
ab00
ab01
ab02
ab03
ab04
ab05
ab06
ab07
ab08
ab09
ab10
ab11
ab12
ab13
ab14
ab15
ab16
ab17
ab18
ab19
ab20
ab21
ab22
ab23
ab24
ab25
ab26
ab27
ab28
ab29
ab30
ab31
ab32
ab33
ab34
ab35
ab36
ab37
ab38
ab39
ab40
ab41
ab42
ab43
ab44
ab45
ab46
ab47
ab48
ab49
ab50
ab51
ab52
ab53
ab54
ab55
ab56
ab57
ab58
ab59
ab60
ab61
ab62
ab63
ab64
ab65
ab66
ab67
ab68
ab69
ab70
ab71
ab72
ab73
ab74
ab75
ab76
ab77
ab78
ab79
ab80
ab81
ab82
ab83
ab84
ab85
ab86
ab87
ab88
ab89
ab90
ab91
ab92
ab93
ab94
ab95
ab96
ab97
ab98
ab99

```

<!-- grade_id: cell-1a555a1a34d8a5a5, points: 4, grade: true, solution: false -->
```python
### DO NOT MODIFY THIS CELL
### BEGIN HIDDEN TESTS
import types
char, num = string.ascii_lowercase, string.digits
usrname_gen = generate_usernames()
tc = unittest.TestCase()
tc.assertIsInstance(usrname_gen, types.GeneratorType)
tc.assertCountEqual([x for x in usrname_gen], [w + x + y + z for w in num for x in num for y in char for z in char])
### END HIDDEN TESTS
```